---
title: 'Data exploring, wrangling, and formatting'
description: ""
---

## Pre-wrangling exploration

```yaml
type: VideoExercise
key: f20c0dcfb8
xp: 50
```

`@projector_key`
146b85090bb8ab77efbfe45c5c751f5d

---

## Visually examine the outcomes with the exposures

```yaml
type: NormalExercise
key: e50ea375f8
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Inspect graphically all variables of interest individually

```yaml
type: TabExercise
key: aec12c0d13
xp: 100
```



`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


***

```yaml
type: NormalExercise
key: 101226608c
xp: 50
```



`@instructions`


`@hint`


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```

***

```yaml
type: NormalExercise
key: 16fbde9ec8
xp: 50
```



`@instructions`


`@hint`


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```

---

## Tidy cohort data and wrangling into analyzable form.

```yaml
type: VideoExercise
key: 3d338af036
xp: 50
```

`@projector_key`
4e1f8ff56b37d8caee655cf2b0b4639d

---

## Transforming and modifying variables

```yaml
type: VideoExercise
key: bfcfbe9aa2
xp: 50
```

`@projector_key`
5d026dadac109f3540f3c1f59a6f96ea

---

## Reduce number of categories in these datasets

```yaml
type: NormalExercise
key: a721377679
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Comparison between different transformations

```yaml
type: TabExercise
key: ca708dca27
xp: 100
```



`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


***

```yaml
type: NormalExercise
key: 0255ca16ff
xp: 35
```



`@instructions`


`@hint`


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```

***

```yaml
type: NormalExercise
key: 2422f8318a
xp: 35
```



`@instructions`


`@hint`


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```

***

```yaml
type: NormalExercise
key: f4dfabce02
xp: 30
```



`@instructions`


`@hint`


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```
